```python
import time
from autograder_term1week5 import *
_globals = globals()
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden </center>

# <center> Worksheet 1: Arithmetic and functions </center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-7383640690456bdc, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

Output:
```
You've changed the contents of the file; thank you!
Your username has been recognised; thank you!

5 out of 5 marks
5
```

## Question 1(i)

Calculate: 

(a) $4^2+1$;
(b) $4^{2+1}$;
(c) $2/3$ as a decimal;
(d) the largest integer smaller than $381/47$;
(e) $4 \times 2^5$;
(f) $(4 \times 2)^5$.

<!-- grade_id: cell-bcf9158f78bdb3a6, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1ia_answer = 4**2+1
q1ib_answer = 4**(2+1)
q1ic_answer = 2/3
q1id_answer = int(381/47)
q1ie_answer = 4*2**5
q1if_answer = (4*2)**5
print(q1ia_answer)
print(q1ib_answer)
print(q1ic_answer)
print(q1id_answer)
print(q1ie_answer)
print(q1if_answer)
# YOUR CODE HERE
```

Output:
```
17
64
0.6666666666666666
8
128
32768
```

<!-- grade_id: cell-593c9367c3c6115a, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1iia_answer = 372517%1820
q1iib_answer = (1417456+1191761)%1820
q1iic_answer = (1417456*1191761)%1820
start = time.time()
q1iid_answer = builtins.pow(1417456,1191761,1820)
duration = time.time() - start
print(q1iia_answer)
print(q1iib_answer)
print(q1iic_answer)
print(q1iid_answer)

# YOUR CODE HERE
372517%1820
(1417456+1191761)%1820
```

Output:
```
1237
1157
636
976
1157
```

<!-- grade_id: cell-b12f71f8ddaf1387, points: 6, grade: true, solution: false -->
```python
# 6 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1(_globals)
```

Output:
```
Question 1(i)a passed!
Question 1(i)b passed!
Question 1(i)c passed!
Question 1(i)d passed!
Question 1(i)e passed!
Question 1(i)f passed!
6 out of 6 tests passed
6
```

## Question 1(ii)

Calculate:

(a) The residue of $372517$ modulo $1820$.

(b) The residue of $(1417456 + 1191761)$ modulo $1820$.

(c) The residue of $(1417456 \times 1191761)$ modulo $1820$.

(d) The residue of $1417456^{1191761}$ modulo $1820$ (for best results use the <code>pow</code> function with three arguments).

<!-- grade_id: cell-d5ff3aef45b2ad85, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete `this cell
# Run this cell for grading
_globals = globals()
question1_ii(_globals)
```

Output:
```
Question 1(ii)a passed!
Question 1(ii)b passed!
Question 1(ii)c passed!
Question 1(ii)d passed!
5 out of 5 tests passed
5
```

## Question 2(i)

Using the <code>math</code> module, calculate: (a) $\sqrt{50}$; (b) $\cos \pi/7$; (c) $e^{e^2}$; (d) $\ln 58$; (e) $\log_{10} 58$; (f) $\cot^{-1} 2$; (g) $\coth^{-1} 2$.

<!-- grade_id: cell-c725db911b14d4ee, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the math module
# Assign your solutions to the questions above to the below variables.
from math import*
q2ia_answer = sqrt(50)
q2ib_answer = cos(pi/7)
q2ic_answer = e**e**2
q2id_answer = log(58)
q2ie_answer = log10(58)
q2if_answer = atan(1/2)
q2ig_answer = atanh(1/2)
print(q2ia_answer)
print(q2ib_answer)
print(q2ic_answer)
print(q2id_answer)
print(q2ie_answer)
print(q2if_answer)
print(q2ig_answer)

```

Output:
```
7.0710678118654755
0.9009688679024191
1618.1779919126518
4.060443010546419
1.7634279935629373
0.46364760900080615
0.5493061443340549
```

<!-- grade_id: cell-660c3f636dc0be36, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_i(_globals)
```

Output:
```
Question 2(i)a passed!
Question 2(i)b passed!
Question 2(i)c passed!
Question 2(i)d passed!
Question 2(i)e passed!
Question 2(i)f passed!
Question 2(i)g passed!
7 out of 7 tests passed
7
```

## Question 2(ii)

The <code>numpy</code> module contains many of the same functions as <code>math</code> (though they're slightly different, in ways we'll explore). Recalculate (a) - (g) using this module instead.

<!-- grade_id: cell-45ed8f2561c9f720, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the numpy module
# Assign your solutions to the questions above to the below variables.
import numpy
q2iia_answer = numpy.sqrt(50)
q2iib_answer = numpy.cos(pi/7)
q2iic_answer = e**e**2
q2iid_answer = log(58)
q2iie_answer = log10(58)
q2iif_answer = arctan(1/2)
q2iig_answer = arctanh(1/2)
print(q2iia_answer)
print(q2iib_answer)
print(q2iic_answer)
print(q2iid_answer)
print(q2iie_answer)
print(q2iif_answer)
print(q2iig_answer)

```

Output:
```
7.0710678118654755
0.9009688679024191
1618.1779919126518
4.060443010546419
1.7634279935629373
0.46364760900080615
0.5493061443340549
```

<!-- grade_id: cell-6390293366653fbb, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_ii(_globals)
```

Output:
```
Question 2(ii)a passed!
Question 2(ii)b passed!
Question 2(ii)c passed!
Question 2(ii)d passed!
Question 2(ii)e passed!
Question 2(ii)f passed!
Question 2(ii)g passed!
7 out of 7 tests passed
7
```

## Question 3

(a) Define variables $p=6$ and $q=9$, and calculate the quantity
$$x=\left(\frac{q}{2}+\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}+\left(\frac{q}{2}-\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}.$$
Verify Cardano's Theorem, that $x$ is a solution of the cubic
$$x^3=p\,x+q.$$

<!-- grade_id: cell-05998432c3deefe2, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 6
q = 9

x = (q/2+((q/2)**2-(p/3)**3)**0.5)**1/3 + (q/2-((q/2)**2-(p/3)**3)**0.5)**1/3
print(p)
print(q)
print(x)
print(x**3-p*x-q)
```

Output:
```
6
9
3.0
0.0
```

<!-- grade_id: cell-721979068553d977, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(a) passed!")
```

Output:
```
Question 3(a) passed!
```

(b) Repeat for $p=3$ and $q=18$.

<!-- grade_id: cell-972547089f6b9cc6, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 3
q = 18

x = (q/2+((q/2)**2-(p/3)**3)**0.5)**(1/3) + (q/2-((q/2)**2-(p/3)**3)**0.5)**(1/3)
print(p)
print(q)
print(x)
print(x**3-p*x-q)
```

Output:
```
3
18
2.999999999999999
-2.1316282072803006e-14
```

<!-- grade_id: cell-22cd955862519f1c, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(b) passed!")
```

Output:
```
Question 3(b) passed!
```

(c) If you try to repeat for $p=6$ and $q=4$, you may run into a problem. Can you solve it?

<!-- grade_id: cell-d7a2fd2365c71adf, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
import cmath
p = 6
q = 4

x = (q/2+((q/2)**2-(p/3)**3)**0.5)**(1/3) + (q/2-((q/2)**2-(p/3)**3)**0.5)**(1/3)
print(p)
print(q)
print(x)

```

Output:
```
6
4
(2.7320508075688776+0j)
```

<!-- grade_id: cell-69e5c37ef72d2e4e, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(c) passed!")
```

Output:
```
Question 3(c) passed!
```

## Question 4

Experiment with <b>integer division</b>, and answer the following question:

For integers $a$ and $b$, the code

<code>a // b</code>

always returns

1. The least integer greater than or equal to $a/b$;
2. The greatest integer less than or equal to $a/b$;
3. The nearest integer to $a/b$;
4. The number $a/b$, rounded towards zero;
5. The number $a/b$, rounded away from zero;
6. None of these; it depends on the signs of $a$ and $b$.

<!-- grade_id: cell-87d8b531ab5cfff5, points: 0, grade: false, solution: true -->
```python
q4_answer = 2
# Assign your solutions to the questions above to the below variables.
45//2
```

Output:
```
22
```

<!-- grade_id: cell-2c3657af708f4b51, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q4_answer == question4())
print("Question 4 passed!")
```

Output:
```
Question 4 passed!
```

## Question 5 (i)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can all be given sensible mathematical values. Using functions from a suitable Python module, calculate these values.

<!-- grade_id: cell-8de18ecb33e6d43d, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
import cmath
q5ia_answer =cmath.asin(-2)
q5ib_answer =cmath.acos(-2)
q5ic_answer =cmath.atanh(-2)
q5id_answer =cmath.log(-2)

print(q5ia_answer)
print(q5ib_answer)
print(q5ic_answer)
print(q5id_answer)


```

Output:
```
(-1.5707963267948966+1.3169578969248166j)
(3.141592653589793-1.3169578969248166j)
(-0.5493061443340549+1.5707963267948966j)
(0.6931471805599453+3.141592653589793j)
```

<!-- grade_id: cell-68368eae35986782, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5_i(_globals)
```

Output:
```
Question 5(i)a passed!
Question 5(i)b passed!
Question 5(i)c passed!
Question 5(i)d passed!
4 out of 4 tests passed
4
```

## Question 5 (ii)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can actually all be calculated using only functions from the <code>numpy</code> module, though there's a bit of a trick to making it work. Find the trick, and make it work!

<!-- grade_id: cell-966f608a75b6df3f, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
import numpy
q5iia_answer =arcsin(-2+0j)
q5iib_answer =arccos(-2+0j)
q5iic_answer =arctanh(-2+0j)
q5iid_answer =log(-2+0j)
print(q5iia_answer)
print(q5iib_answer)
print(q5iic_answer)
print(q5iid_answer)


```

Output:
```
(-1.5707963267948966+1.3169578969248166j)
(3.141592653589793-1.3169578969248166j)
(-0.5493061443340549+1.5707963267948966j)
(0.6931471805599453+3.141592653589793j)
```

<!-- grade_id: cell-c794ccf9d2ca10bb, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
question5_ii(_globals)
```

Output:
```
Question 5(ii)a passed!
Question 5(ii)b passed!
Question 5(ii)c passed!
Question 5(ii)d passed!
4 out of 4 tests passed
4
```

## Question 5 (iii)

(a) Using the built-in `pow` function, calculate the residue of $7725423^{3835624}$ modulo $25231211$, assigning your answer to the variable `q5iiia_answer`.

<!-- grade_id: cell-9c8b99fcd38bc268, points: 0, grade: false, solution: true -->
```python
# Assign your solution to the questions above to the below variables.
import builtins
question5iiia_answer = builtins.pow(7725423,3835624,25231211)
print(question5iiia_answer)


```

Output:
```
20666298
```

<!-- grade_id: cell-e00764f878a596d0, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiia_answer == question5_iii_a())
print('test case passed!')
```

Output:
```
test case passed!
```

(b) Wild-card import all the functions from the `math` module as follows:

```python
from math import *
```

Then, attempt your calculation in part (a) for a second time, and note what happens. 

Explore a bit further if you like. How would you explain what has happened?

<ol>
    <li>Wild-card importing all the functions from the <code>math</code> module breaks the <code>pow</code> function; no function called <code>pow</code> is now available.</li>
    <li>Wild-card importing all the functions from the <code>math</code> module produces a namespace clash between the function called <code>pow</code> that belongs to that module and the built-in function of the same name; from now on, only the <code>math</code> version works, and that can't be used to calculate modular powers.</li>
    <li>Wild-card importing all the functions from the <code>math</code> module produces a namespace clash between the function called <code>pow</code> that belongs to that module and the built-in function of the same name; from now on, only the built-in version works, and that can't be used to calculate modular powers.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-51434c057025b9dd, points: 0, grade: false, solution: true -->
```python
question5iiib_answer = 2

```

<!-- grade_id: cell-b0a3a702d0ea9cfd, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiib_answer == question5_iii_b())
print('test case passed!')
```

Output:
```
test case passed!
```

(c) Try the following lines of code:

```python
from math import *
import builtins
builtins.pow(7725423,3835624,25231211)
```

Which of the following statements is the closest to the truth?

<ol>
    <li>Once a built-in function has been overwritten in a namespace clash, it can't be used again unless the session is restarted.</li>
    <li>Namespace clashes with built-in functions can be rectified by importing the <code>builtins</code> module, meaning that there's no need to worry about them.</li>
    <li>Namespace clashes with built-in functions can be rectified by importing the <code>builtins</code> module, so they may not be fatal, but they're still worth taking care to avoid.</li>
</ol>

(This may be a bit of a matter of opinion; but our opinion on it is strong, and we bet you can guess it.)

<!-- grade_id: cell-786585eb3ddbb615, points: 0, grade: false, solution: true -->
```python
question5iiic_answer = 3
# YOUR CODE HERE

```

<!-- grade_id: cell-f1e70db763174c19, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiic_answer == question5_iii_c())
print('test case passed!')
```

Output:
```
test case passed!
```

```python

```