```python
import time
from autograder_term1week5 import *
_globals = globals()
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden </center>

# <center> Worksheet 1: Arithmetic and functions </center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-7383640690456bdc, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

## Question 1(i)

Calculate: 

(a) $4^2+1$;
(b) $4^{2+1}$;
(c) $2/3$ as a decimal;
(d) the largest integer smaller than $381/47$;
(e) $4 \times 2^5$;
(f) $(4 \times 2)^5$.

<!-- grade_id: cell-bcf9158f78bdb3a6, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1ia_answer = 
q1ib_answer = 
q1ic_answer = 
q1id_answer = 
q1ie_answer = 
q1if_answer = 
print(q1ia_answer)
print(q1ib_answer)
print(q1ic_answer)
print(q1id_answer)
print(q1ie_answer)
print(q1if_answer)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b12f71f8ddaf1387, points: 6, grade: true, solution: false -->
```python
# 6 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1(_globals)
```

## Question 1(ii)

Calculate:

(a) The residue of $372517$ modulo $1820$.

(b) The residue of $(1417456 + 1191761)$ modulo $1820$.

(c) The residue of $(1417456 \times 1191761)$ modulo $1820$.

(d) The residue of $1417456^{1191761}$ modulo $1820$ (for best results use the <code>pow</code> function with three arguments).

<!-- grade_id: cell-593c9367c3c6115a, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1iia_answer = 
q1iib_answer = 
q1iic_answer = 
start = time.time()
q1iid_answer = 
duration = time.time() - start
print(q1iia_answer)
print(q1iib_answer)
print(q1iic_answer)
print(q1iid_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d5ff3aef45b2ad85, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete `this cell
# Run this cell for grading
_globals = globals()
question1_ii(_globals)
```

## Question 2(i)

Using the <code>math</code> module, calculate: (a) $\sqrt{50}$; (b) $\cos \pi/7$; (c) $e^{e^2}$; (d) $\ln 58$; (e) $\log_{10} 58$; (f) $\cot^{-1} 2$; (g) $\coth^{-1} 2$.

<!-- grade_id: cell-c725db911b14d4ee, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the math module
# Assign your solutions to the questions above to the below variables.
q2ia_answer = 
q2ib_answer = 
q2ic_answer = 
q2id_answer = 
q2ie_answer = 
q2if_answer = 
q2ig_answer = 
print(q2ia_answer)
print(q2ib_answer)
print(q2ic_answer)
print(q2id_answer)
print(q2ie_answer)
print(q2if_answer)
print(q2ig_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-660c3f636dc0be36, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_i(_globals)
```

## Question 2(ii)

The <code>numpy</code> module contains many of the same functions as <code>math</code> (though they're slightly different, in ways we'll explore). Recalculate (a) - (g) using this module instead.

<!-- grade_id: cell-45ed8f2561c9f720, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the numpy module
# Assign your solutions to the questions above to the below variables.
q2iia_answer = 
q2iib_answer = 
q2iic_answer = 
q2iid_answer = 
q2iie_answer = 
q2iif_answer = 
q2iig_answer = 
print(q2iia_answer)
print(q2iib_answer)
print(q2iic_answer)
print(q2iid_answer)
print(q2iie_answer)
print(q2iif_answer)
print(q2iig_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6390293366653fbb, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_ii(_globals)
```

## Question 3

(a) Define variables $p=6$ and $q=9$, and calculate the quantity
$$x=\left(\frac{q}{2}+\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}+\left(\frac{q}{2}-\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}.$$
Verify Cardano's Theorem, that $x$ is a solution of the cubic
$$x^3=p\,x+q.$$

<!-- grade_id: cell-05998432c3deefe2, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 
q = 

x = 
print(p)
print(q)
print(x)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-721979068553d977, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(a) passed!")
```

(b) Repeat for $p=3$ and $q=18$.

<!-- grade_id: cell-972547089f6b9cc6, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 
q = 

x = 
print(p)
print(q)
print(x)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-22cd955862519f1c, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(b) passed!")
```

(c) If you try to repeat for $p=6$ and $q=4$, you may run into a problem. Can you solve it?

<!-- grade_id: cell-d7a2fd2365c71adf, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 
q = 

x = 
print(p)
print(q)
print(x)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69e5c37ef72d2e4e, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(c) passed!")
```

## Question 4

Experiment with <b>integer division</b>, and answer the following question:

For integers $a$ and $b$, the code

<code>a // b</code>

always returns

1. The least integer greater than or equal to $a/b$;
2. The greatest integer less than or equal to $a/b$;
3. The nearest integer to $a/b$;
4. The number $a/b$, rounded towards zero;
5. The number $a/b$, rounded away from zero;
6. None of these; it depends on the signs of $a$ and $b$.

<!-- grade_id: cell-87d8b531ab5cfff5, points: 0, grade: false, solution: true -->
```python
q4_answer = # enter one of 1, 2, 3, 4, 5 or 6
# Assign your solutions to the questions above to the below variables.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c3657af708f4b51, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q4_answer == question4())
print("Question 4 passed!")
```

## Question 5 (i)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can all be given sensible mathematical values. Using functions from a suitable Python module, calculate these values.

<!-- grade_id: cell-8de18ecb33e6d43d, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.

q5ia_answer =
q5ib_answer =
q5ic_answer =
q5id_answer =

print(q5ia_answer)
print(q5ib_answer)
print(q5ic_answer)
print(q5id_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-68368eae35986782, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5_i(_globals)
```

## Question 5 (ii)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can actually all be calculated using only functions from the <code>numpy</code> module, though there's a bit of a trick to making it work. Find the trick, and make it work!

<!-- grade_id: cell-966f608a75b6df3f, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q5iia_answer =
q5iib_answer =
q5iic_answer =
q5iid_answer =
print(q5iia_answer)
print(q5iib_answer)
print(q5iic_answer)
print(q5iid_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c794ccf9d2ca10bb, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
question5_ii(_globals)
```

## Question 5 (iii)

(a) Using the built-in `pow` function, calculate the residue of $7725423^{3835624}$ modulo $25231211$, assigning your answer to the variable `q5iiia_answer`.

<!-- grade_id: cell-9c8b99fcd38bc268, points: 0, grade: false, solution: true -->
```python
# Assign your solution to the questions above to the below variables.
question5iiia_answer = 
print(q5iiia_answer)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e00764f878a596d0, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiia_answer == question5_iii_a())
print('test case passed!')
```

(b) Wild-card import all the functions from the `math` module as follows:

```python
from math import *
```

Then, attempt your calculation in part (a) for a second time, and note what happens. 

Explore a bit further if you like. How would you explain what has happened?

<ol>
    <li>Wild-card importing all the functions from the <code>math</code> module breaks the <code>pow</code> function; no function called <code>pow</code> is now available.</li>
    <li>Wild-card importing all the functions from the <code>math</code> module produces a namespace clash between the function called <code>pow</code> that belongs to that module and the built-in function of the same name; from now on, only the <code>math</code> version works, and that can't be used to calculate modular powers.</li>
    <li>Wild-card importing all the functions from the <code>math</code> module produces a namespace clash between the function called <code>pow</code> that belongs to that module and the built-in function of the same name; from now on, only the built-in version works, and that can't be used to calculate modular powers.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-51434c057025b9dd, points: 0, grade: false, solution: true -->
```python
question5iiib_answer = # (delete, and fill in the value 1, 2 or 3)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b0a3a702d0ea9cfd, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiib_answer == question5_iii_b())
print('test case passed!')
```

(c) Try the following lines of code:

```python
from math import *
import builtins
builtins.pow(7725423,3835624,25231211)
```

Which of the following statements is the closest to the truth?

<ol>
    <li>Once a built-in function has been overwritten in a namespace clash, it can't be used again unless the session is restarted.</li>
    <li>Namespace clashes with built-in functions can be rectified by importing the <code>builtins</code> module, meaning that there's no need to worry about them.</li>
    <li>Namespace clashes with built-in functions can be rectified by importing the <code>builtins</code> module, so they may not be fatal, but they're still worth taking care to avoid.</li>
</ol>

(This may be a bit of a matter of opinion; but our opinion on it is strong, and we bet you can guess it.)

<!-- grade_id: cell-786585eb3ddbb615, points: 0, grade: false, solution: true -->
```python
question5iiic_answer = # (delete, and fill in the value 1, 2 or 3)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1e70db763174c19, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5iiic_answer == question5_iii_c())
print('test case passed!')
```

```python

```